Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Fourier Transform 2021

```python
import numpy as np


def sinc(x):
    return np.sinc(x / np.pi)
```

Note:  if you need to use the sinc function, please use the one provided above rather the one in the NumPy library.

## Part A

Find the Fourier transform for $f(t)$ below.

<img src="fourier_transform2021a.png" width="400"/>

Represent your answer as a function `F(omega)`.

<!-- grade_id: cell-9c81bd6cbb6aa4b7, points: 0, grade: false, solution: true -->
```python
def F(omega):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-019622881b2429a2, points: 2, grade: true, solution: false -->
```python
assert isinstance(F(0), (float, complex))


def f(t):
    t = np.asarray(t)
    result = np.zeros_like(t, dtype=float)
    mask1 = (-2 <= t) & (t < -1)
    result[mask1] = t[mask1] + 2
    mask2 = (-1 <= t) & (t < 0)
    result[mask2] = -t[mask2]
    mask3 = (0 <= t) & (t < 1)
    result[mask3] = t[mask3]
    mask4 = (1 <= t) & (t < 2)
    result[mask4] = 2 - t[mask4]
    return result if result.shape != () else float(result)


t = np.linspace(-5, 5, 1000)

dt = t[1] - t[0]
f_t = f(t)
F_num = np.fft.fftshift(np.fft.fft(f_t, n=2**14)) * dt
omega_num = np.fft.fftshift(np.fft.fftfreq(len(F_num), dt)) * 2 * np.pi

omega_vals = np.linspace(-10, 10, 100)
F_vals = np.array([F(om) for om in omega_vals])


from matplotlib import pyplot as plt

plt.figure(figsize=(10, 4))
plt.subplot(1, 2, 1)
plt.plot(t, f_t)
plt.xlabel(r"$t$")
plt.ylabel(r"$f(t)$")
plt.title("Function f(t)")
plt.grid()

plt.subplot(1, 2, 2)
plt.plot(omega_num, np.abs(F_num), label="Numerical FFT")
plt.plot(omega_vals, np.abs(F_vals), "--", label="Analytical F(ω)")
plt.legend()
plt.xlim(-10, 10)
plt.xlabel(r"$\omega$")
plt.ylabel(r"$|F(\omega)|$")
plt.title("Magnitude of the Fourier Transform of f(t)")
plt.grid()
plt.show()


```

## Part B

Find the inverse transform $f(t)$ if $|F(\omega)|$ and $\angle F(\omega)$ are given below.  

<img src="fourier_transform2021b.png" width="800"/>

Represent your answer as a function `f(t)`.  

<!-- grade_id: cell-3b6f47a4fd843f2b, points: 0, grade: false, solution: true -->
```python
def f(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e8049c8142664a0e, points: 2, grade: true, solution: false -->
```python
assert isinstance(f(0), (float, complex))


def F(omega):
    F_omega = np.zeros_like(omega, dtype=complex)
    mask1 = (0 <= omega) & (omega < 1)
    mask2 = (2 <= omega) & (omega < 3)
    F_omega[mask1] = 1.0
    F_omega[mask2] = 1.0
    return F_omega


# Define the frequency-domain signal F(omega)
d_omega = 0.01
omega = np.arange(-40, 40, d_omega)
F_omega = F(omega)

# Compute the numerical inverse Fourier Transform
# The scaling factor is necessary to approximate the continuous integral
n = len(omega)
scaling_factor = n * d_omega / (2 * np.pi)
f_t_numerical = scaling_factor * np.fft.ifftshift(
    np.fft.ifft(np.fft.ifftshift(F_omega))
)

# Create the corresponding time vector for the output
t = np.fft.fftshift(np.fft.fftfreq(n, d=d_omega / (2 * np.pi)))


f_t_user = f(t)

# Plot the results for comparison
plt.figure(figsize=(12, 6))
plt.subplot(1, 2, 1)
plt.plot(omega, np.abs(F_omega), label="Numerical F(ω)", linewidth=4)
plt.xlim(-4, 4)
plt.xlabel("ω")
plt.ylabel("|F(ω)|")
plt.title("Absolute Value of F(ω)")


plt.subplot(1, 2, 2)
plt.plot(t, np.real(f_t_numerical), label="Numerical Inverse FT, real", linewidth=4)
plt.plot(t, np.imag(f_t_numerical), label="Numerical Inverse FT, imag", linestyle="--")
plt.plot(t, f_t_user.real, label="User's Analytical f(t), real", linestyle=":")
plt.plot(t, f_t_user.imag, label="User's Analytical f(t), imag", linestyle="-.")
plt.xlim(-5, 10)
plt.xlabel("t")
plt.ylabel("f(t)")
plt.title("Comparison of Numerical and Analytical Inverse Fourier Transforms")
plt.legend()
plt.grid(True)
plt.show()


```