Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Fourier transform 2024

```python
from matplotlib import pyplot as plt
import numpy as np


def sinc(x):
    return np.sinc(x / np.pi)
```

## Part A

Consider the system described by an input/output pair.

<img src="fourier_transform2024a.png" width="800"/>

Find $Y(\omega)$, the Fourier transform of the output $y(t)$.  Feel free to use Fourier properties and known transforms or brute-force integration.  Simplify your answer as much as possible.  Represent your answer as a function `Y(omega)`.  Hint: to sanity check your solution, look at the symmetry of your final answer and make sure it makes sense!

<!-- grade_id: cell-6a7a8f2ced884e40, points: 0, grade: false, solution: true -->
```python
def Y(omega):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a44409b6209694cf, points: 2, grade: true, solution: false -->
```python
assert isinstance(Y(1), (float, complex))

y = np.vectorize(lambda t: t if -1 <= t < 1 else 0.0)

t = np.linspace(-5, 5, 1000)

dt = t[1] - t[0]
y_t = y(t)
Y_num = np.fft.fftshift(np.fft.fft(y_t, n=2**14)) * dt
omega_num = np.fft.fftshift(np.fft.fftfreq(len(Y_num), dt)) * 2 * np.pi

omega_vals = np.linspace(-10, 10, 100)
Y_vals = np.array([Y(om) for om in omega_vals])


from matplotlib import pyplot as plt

plt.figure(figsize=(10, 4))
plt.subplot(1, 2, 1)
plt.plot(t, y_t)
plt.xlabel(r"$t$")
plt.ylabel(r"$y(t)$")
plt.title("Function y(t)")
plt.grid()

plt.subplot(1, 2, 2)
plt.plot(omega_num, np.abs(Y_num), label="Numerical FFT")
plt.plot(omega_vals, np.abs(Y_vals), "--", label="Analytical Y(ω)")
plt.legend()
plt.xlim(-10, 10)
plt.xlabel(r"$\omega$")
plt.ylabel(r"$|Y(\omega)|$")
plt.title("Magnitude of the Fourier Transform of y(t)")
plt.grid()
plt.show()


omega = np.linspace(-10, 10, 1000)

plt.plot(omega, np.real(Y(omega)), label="Real")
plt.plot(omega, np.imag(Y(omega)), label="Imaginary")
plt.grid()
plt.legend()
plt.xlabel("Frequency (ω)")
plt.ylabel("Y(ω)")
plt.title("Real and Imaginary Parts of Y(ω)")

fig, ax1 = plt.subplots()


# Magnitude on left y-axis
ax1.plot(omega, np.abs(Y(omega)), "b-", label="Magnitude")
ax1.set_xlabel("Frequency (ω)")
ax1.set_ylabel("Magnitude", color="b")
ax1.tick_params(axis="y", labelcolor="b")
ax1.grid()

# Phase on right y-axis
ax2 = ax1.twinx()
ax2.plot(omega, np.angle(Y(omega)), "r-", label="Phase")
ax2.set_ylabel("Phase [rad/s]", color="r")
ax2.tick_params(axis="y", labelcolor="r")

plt.title("Magnitude and Phase of Y(ω)")


```

Given this input/output relationship, is this system linear and time-invariant?  Please respond with `ans = "LTI"` or `ans = "not LTI"`.

<!-- grade_id: cell-4f0904b7cfd600ee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b5b66d482c426fff, points: 1, grade: true, solution: false -->
```python
assert ans in ["LTI", "not LTI"]

```

## Part B

Let $x(t)$ be a signal whose spectrum is identically zero outside the range $-\omega_0 \leq \omega \leq \omega_0$. An example of such a spectrum is shown below.  The actual values of $X(\omega)$ are not important for this problem, but the range of $\omega$ for which $X(\omega)$ is non-zero is.

<img src="fourier_transform2024b.png" width="500"/>

For the following signals, determine the range of $\omega$ (in terms of $\omega_0$ over which their spectrum is non-zero.


$$y(t)=x(t)+x(t-1)$$

Represent your answer as a function `non_zero1(omega, omega_0)` that indicates `True` if $Y(\omega)$ is non-zero at some $\omega$ given some $\omega_0$ or `False` if $Y(\omega)$ is zero at some $\omega$ given some $\omega_0$.

<!-- grade_id: cell-bbdd060cd7298f6a, points: 0, grade: false, solution: true -->
```python
def non_zero1(omega, omega_0):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3273ac6ad2abe4b3, points: 1, grade: true, solution: false -->
```python
assert isinstance(non_zero1(1, 2), bool)


```

$$y(t) = x(t) \cos(\omega_0 t)$$

Represent your answer as a function `non_zero2(omega, omega_0)` that indicates `True` if $Y(\omega)$ is non-zero at some $\omega$ given some $\omega_0$ or `False` if $Y(\omega)$ is zero at some $\omega$ given some $\omega_0$.

<!-- grade_id: cell-39240ff61f1d0b82, points: 0, grade: false, solution: true -->
```python
def non_zero2(omega, omega_0):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-64af5dc98caa9045, points: 1, grade: true, solution: false -->
```python
assert isinstance(non_zero2(1, 2), bool)


```