Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Fourier Tranform 2019

```python
import numpy as np


def sinc(x):
    return np.sinc(x / np.pi)
```

Note:  if you need to use the sinc function, please use the one provided above rather the one in the NumPy library.

Find the Fourier transform of $f(t) = 2\text{sinc}(\pi t) \cos(\pi t)$.  You may use tables or direct integration.

Represent your answer as a a function `F(omega)`.

<!-- grade_id: cell-92b4b8c216dd9b10, points: 0, grade: false, solution: true -->
```python
def F(omega):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9c21ca5d47deaf20, points: 3, grade: true, solution: false -->
```python
assert isinstance(F(0), (float, complex))


def f(t):
    return 2 * sinc(np.pi * t) * np.cos(np.pi * t)


t = np.linspace(-10, 10, 1000)

dt = t[1] - t[0]
f_t = f(t)
F_num = np.fft.fftshift(np.fft.fft(f_t, n=2**14)) * dt
omega_num = np.fft.fftshift(np.fft.fftfreq(len(F_num), dt)) * 2 * np.pi

omega_vals = np.linspace(-10, 10, 100)
F_vals = np.array([F(om) for om in omega_vals])


from matplotlib import pyplot as plt

plt.figure(figsize=(10, 4))
plt.subplot(1, 2, 1)
plt.plot(t, f_t)
plt.xlabel(r"$t$")
plt.ylabel(r"$f(t)$")
plt.title("Function f(t)")
plt.grid()

plt.subplot(1, 2, 2)
plt.plot(omega_num, np.abs(F_num), label="Numerical FFT")
plt.plot(omega_vals, np.abs(F_vals), "--", label="Analytical F(ω)")
plt.legend()
plt.xlim(-10, 10)
plt.xlabel(r"$\omega$")
plt.ylabel(r"$|F(\omega)|$")
plt.title("Magnitude of the Fourier Transform of f(t)")
plt.grid()
plt.show()


```