Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Parseval's Theorem

```python
import numpy as np
```

Parseval's theorem states that
$$
\int_{-\infty}^{\infty} |x(t)|^2 dt = \frac{1}{2\pi} \int_{-\infty}^{\infty} |X(\omega)|^2 d\omega.
$$
In this equation, $x(t)$ is the signal in the time domain and $X(\omega)$ is the signal in the frequency domain. The theorem states that the total energy of a signal in time domain is equal to the total energy of the signal in the frequency domain.

The spectrum
$$
Y(\omega) = \frac{1}{1-j\omega}
$$
has an energy
$$
E_Y = \frac{1}{2\pi}\int^\infty_{-\infty} \left|\frac{1}{1-j\omega} \right|^2 d\omega.
$$
Find $E_Y$.  Represent your answer as `Ey = XX`.

<!-- grade_id: cell-bff2960fa8f8134e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b53c8aaef0eb5f83, points: 2, grade: true, solution: false -->
```python
assert isinstance(Ey, float)

```

We can extend this to a _generalized_ Parseval's theorem, which is
$$
\int_{-\infty}^{\infty} x_1(t) x_2^*(t) dt = \frac{1}{2\pi} \int_{-\infty}^{\infty} X_1(\omega) X_2^*(\omega) d\omega
$$

In this equation, $x_1(t)$ and $x_2(t)$ are signals in the time domain and $X_1(\omega)$ and $X_2(\omega)$ are their respective signals in the frequency domain. The theorem states that the integral of the product of a function and its conjugate in the time domain is equal to the integral of the product of their Fourier transforms in the frequency domain.

Among other useful attributes, we can use Parseval's theorem to compute some really weird integrals.

Please solve the integral:
$$
\int^\infty_{-\infty} \cos(2\pi t) \text{sinc}^2(t) \text{sinc}(2t)dt
$$
Represent your answer as `ans = XX`.

<!-- grade_id: cell-9ed5cadaacbafbe6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d16de7a14b7ac10, points: 2, grade: true, solution: false -->
```python
assert isinstance(ans, float)


```