Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Quadrature Modulation

```python
from matplotlib import pyplot as plt
import numpy as np
```

Let's look at a quadrature amplitude modulation system.  The block diagram of the modulation system is 

<img src="quad_modulation.png" width="500"/>  


Assume that the inputs are real signals.  Note that the $j$ factor for the lower modulator.



```python
def x(t):
    return np.cos(t) + np.cos(2 * t)


def y(t):
    return np.sin(t) + np.sin(2 * t)
```

<!-- grade_id: cell-0332ebbb73945d41, points: 0, grade: false, solution: true -->
```python
def z(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-52840e58893db6ac, points: 2, grade: true, solution: false -->
```python
assert isinstance(z(0), complex)


```

Find the spectrum of the transmitted signal $Z(\omega)$. Simplify your expression, and eliminate any convolutions.  You are given two functions for $X(\omega)$ and $Y(\omega)$ below.

Represent your answer as a function `Z(omega)`.

```python
def X(omega):
    if -np.pi <= omega <= np.pi:
        return 1.0
    else:
        return 0.0


def Y(omega):
    if -np.pi <= omega <= np.pi:
        return 1 - abs(omega) / np.pi
    else:
        return 0.0


omega_vals = np.linspace(-4 * np.pi, 4 * np.pi, 1000)
X_vals = np.array([X(omega) for omega in omega_vals])
Y_vals = np.array([Y(omega) for omega in omega_vals])

plt.figure()
plt.plot(omega_vals, X_vals, label=r"$X(\omega)$")
plt.plot(omega_vals, Y_vals, label=r"$Y(\omega)$")
plt.xlabel(r"$\omega$")
plt.legend()
plt.grid()
```

<!-- grade_id: cell-81f8c68f32000387, points: 0, grade: false, solution: true -->
```python
def Z(omega):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9e6b3d38e44830e2, points: 2, grade: true, solution: false -->
```python
assert isinstance(Z(0), (float, complex))

z_vals = np.array([Z(omega) for omega in omega_vals])

plt.figure()
plt.plot(omega_vals, z_vals, label=r"$Z(\omega)$")
plt.xlabel(r"$\omega$")
plt.legend()
plt.grid()


```