Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# System Response 2022

```python
from matplotlib import pyplot as plt
import numpy as np
```

The following function does not need to be completed.  You are welcome to complete it and use it to check your work in the plots below.  You will not be graded on this code.  If you don't want to use this, just have it return zeros or something so the code below won't break.

<!-- grade_id: cell-66d35ab6c30a6ed9, points: 0, grade: false, solution: true -->
```python
def fourier_series(D_n, omega_0, t):
    """Calculates the Fourier series of a periodic function with a given set of
    Fourier coefficients D_n, angular frequency omega, and time t.

    Parameters
    ----------
    D_n : ndarray
        An array of Fourier coefficients from -N <= n <= N
    omega_0 : float
        The angular frequency of the periodic function
    t : ndarray
        An array of time values

    Returns
    -------
    ndarray
        An array of function values at each time value.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Determine the period and fundamental frequency $\omega_0$ for the function $f(t)$ seen below.

<img src="system_response2022a.png" width="400"/>  

Represent your answer as `omega_0 = XX`.

<!-- grade_id: cell-83da82e261fa969d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-138058510381cecf, points: 1, grade: true, solution: false -->
```python
assert isinstance(omega_0, float)

### BEGIN HIIDDEN TESTS
assert np.isclose(omega_0, np.pi / 2)
```

Derive the complex exponential Fourier series for $f(t)$.  Report your values of $D_n$ as a function `D(n)`.

<!-- grade_id: cell-0719085b253213f7, points: 0, grade: false, solution: true -->
```python
def D(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Let's check to see if this matches the original $f(t)$ using $N=50$.

<!-- grade_id: cell-717ac010545e2499, points: 2, grade: true, solution: false -->
```python
assert isinstance(D(1), (float, complex))

D_n = np.array([D(n) for n in range(-50, 51)])

t = np.linspace(-4, 4, 1000)
f_t = fourier_series(D_n, omega_0, t).real

plt.figure()
plt.plot(t, f_t, label="f(t)")
plt.xlabel("time")
plt.ylabel("$f(t)$")
plt.title("Fourier Series of $f(t)$")
plt.grid()


```

Derive the transfer function $H(s)$ for the circuit seen below.  

<img src="system_response2022b.png" width="400"/>  

Represent your answer as a function `H(s)`.

<!-- grade_id: cell-98ed30ab7debaf50, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b707900e5b22a35f, points: 1, grade: true, solution: false -->
```python
assert isinstance(H(1), float)


```

Given the input $f(t)$ for the circuit above, what is the output $y(t)$?

Represent your answer as a function `y(t)`.  Make sure that in your summation you are only summing terms $-50 \leq n \leq 50$.

<!-- grade_id: cell-c6b7ccdf18b94c5c, points: 0, grade: false, solution: true -->
```python
def y(t):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Let's visualize this result and see if it makes sense.  Does this behavior make sense intuitively if a square wave goes into a $RL$ circuit?

<!-- grade_id: cell-f0871eb6638f36c3, points: 2, grade: true, solution: false -->
```python
assert isinstance(D(1), (float, complex))


t = np.linspace(-4, 4, 1000)
y_t = y(t).real

plt.figure()
plt.plot(t, f_t, label="f(t)")
plt.plot(t, y_t, label="y(t)")
plt.xlabel("time")
plt.ylabel("$f(t)$")
plt.title("Fourier Series of $f(t)$")
plt.legend()
plt.grid()


```